---
title: 'CHAPTER 2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 41e0a0d2e7
xp: 100
```

ABCABCABCABCABCABCABCABCABCABCABCABCABC
Print "Hi"

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
print("Hi!")
```

`@sct`
```{r}

```
