---
title: 'CHAPTER 2'
description: hjjh
---

## Insert exercise title here

```yaml
type: NormalExercise
key: a0aae578d5
xp: 100
```

dewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcewdewfcew

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
